Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "swafdilck"
COLLABORATORS = "swafdilck,et.al"
```

---

# Module 2 Assessed Exercise: More Manipulations with the Field Calculator

This template aims to help you developing the key functions for your MGRS coodinates translator. Specificly, you will need to develop two functions for the simple local case: one for translating x coordinates, one for translating y coorinates. Similarly, two more additional cells will be given for your more generic implementation. Here, we won't be using pyqgis. The functions you developed here can be, however, used in the Field Calculator in QGIS GUI for UTM coordinates calculation.

### Local MGRS/UTM translator (i.e., the 48R Grid Zone)
Please use the following cells for developing codes for the simple case. (Reminder: your code should be able to handle spaces within MGRS string as shown in the example dataset. This is a common task when you work with human input data). Example input: 'YS 374671'

```python
# !pip install mgrs
```

```python
import mgrs
```

```python
Input=input()
with open('coordinate.txt', 'w') as f:
   f.write(Input)
```

Output:
```
48RYS374671
```

```python
def mgrsToUtmX(mgrs_raw):
    '''return the UTM x coordinate as a number of a given MGRS coordinates string'''

    import mgrs

    try:

     file=open("coordinate.txt","r+",encoding="utf-8") 
     coordinate=file.readlines() 

    except FileNotFoundError:
      print("coordinate.txt not found!")
      exit()

    file.seek(0) #The top row operation in the coordinate.txt file

     #Determining whether the coordinates are MGRS or Geographic.
    firstline=file.readline()
    spaceindex=firstline.find(" ")
    control=firstline[spaceindex]

    if control == " ":
     print ("Geographic > MGRS conversion type is selected.")
     for c in coordinate:
      
        m = mgrs.MGRS()
        c=c.encode() #Converting string variable to byte data type
        d = m.MGRSToUTM(c)
        d=str(d)
        index=d.find(",")
        l=len(d)
        latitude=d[1:index]
        longitude=d[(index+2):(l-1)]
        file2=open("coordinate_UTM.txt","a+",encoding="utf-8")
        file2.write(longitude+" "+latitude+"\n")

(latitude)
```

Output:
```
'39.940574445115615'
```

```python
convert = mgrs.MGRS()
coords = u"48RYS374671"
convert.MGRSToUTM(coords)
```

Output:
```
(48, 'N', 737400.0, 3167100.0)
```

<!-- grade_id: local_mgrsx_test1, points: 1, grade: true, solution: false -->
```python
'''test your implementation with example test case'''
assert mgrsToUtmX('YS374671')==737400, 'YS374671 should be 737400 in UTM x coordinate'
print('success!')
```

<!-- grade_id: cell-7449bc8ff985add3, points: 1, grade: true, solution: false -->
```python
'''hidden tests will be executed after submission'''
```

<!-- grade_id: local_mgrsy, points: 0, grade: false, solution: true -->
```python
def mgrsToUtmY(mgrs_raw):
    '''return the UTM y coordinate as a number of a given MGRS coordinates string'''
    import mgrs

    try:

     file=open("coordinate.txt","r+",encoding="utf-8") 
     coordinate=file.readlines() 

    except FileNotFoundError:
      print("coordinate.txt not found!")
      exit()

    file.seek(0) #The top row operation in the coordinate.txt file

     #Determining whether the coordinates are MGRS or Geographic.
    firstline=file.readline()
    spaceindex=firstline.find(" ")
    control=firstline[spaceindex]

    if control == " ":
     print ("Geographic > MGRS conversion type is selected.")
     for c in coordinate:
      
        m = mgrs.MGRS()
        c=c.encode() #Converting string variable to byte data type
        d = m.MGRSToUTM(c)
        d=str(d)
        index=d.find(",")
        l=len(d)
        latitude=d[1:index]
        longitude=d[(index+2):(l-1)]
        file2=open("coordinate_UTM.txt","a+",encoding="utf-8")
        file2.write(longitude+" "+latitude+"\n")

(longitude)

```

Output:
```
'32.852575679878804'
```

<!-- grade_id: local_mgrsy_test1, points: 1, grade: true, solution: false -->
```python
'''test your implementation with an example test case'''
assert mgrsToUtmY('YS374671')==1167100, 'YS374671 should be 1167100 in UTM y coordinate'
print('success!')
```

<!-- grade_id: cell-36a97c19da14ea2a, points: 1, grade: true, solution: false -->
```python
'''hidden tests will be executed after submission'''
```

### Global MGRS/UTM translator
In this section, you will extend the functionality of the translaor you just implemented allowing it to be able to deal with data from different Grid Junction Zones. The full MGRS coordinates including its zone code in the front will be provided as input to the functions (e.g., '31U CT0376087415').

<!-- grade_id: local_mgrsx, points: 0, grade: false, solution: true -->
```python
Input=input()
with open('coordinate.txt', 'w') as f:
   f.write(Input)
```

Output:
```
48RYS374671
```

<!-- grade_id: global_mgrsx, points: 0, grade: false, solution: true -->
```python
def mgrsToUtmX_Global(mgrs_raw):
    '''return the UTM x coordinate as a number of a given MGRS coordinates string'''

    import mgrs

    try:

     file=open("coordinate.txt","r+",encoding="utf-8") 
     coordinate=file.readlines() 

    except FileNotFoundError:
      print("coordinate.txt not found!")
      exit()

    file.seek(0) #The top row operation in the coordinate.txt file

     #Determining whether the coordinates are MGRS or Geographic.
    firstline=file.readline()
    spaceindex=firstline.find(" ")
    control=firstline[spaceindex]

    if control == " ":
     print ("Geographic > MGRS conversion type is selected.")
     for c in coordinate:
      
        m = mgrs.MGRS()
        c=c.encode() #Converting string variable to byte data type
        d = m.MGRSToUTM(c)
        d=str(d)
        index=d.find(",")
        l=len(d)
        latitude=d[1:index]
        longitude=d[(index+2):(l-1)]
        file2=open("coordinate_UTM.txt","a+",encoding="utf-8")
        file2.write(longitude+" "+latitude+"\n")

(latitude)
```

Output:
```
'39.940574445115615'
```

<!-- grade_id: global_mgrsx_tests, points: 1, grade: true, solution: false -->
```python
'''test your implementation with an example test case'''
assert mgrsToUtmX_Global('48RYS374671')==737400, '48RYS374671 should be 737400 in UTM x coordinate'
print('success!')
```

<!-- grade_id: cell-64b630c5832ceaa4, points: 1, grade: true, solution: false -->
```python
'''hidden tests will be executed after submission'''
```

<!-- grade_id: global_mgrsy, points: 0, grade: false, solution: true -->
```python
def mgrsToUtmY_Global(mgrs_raw):
    '''return the UTM y coordinate as a number of a given MGRS coordinates string'''
    import mgrs

    try:

     file=open("coordinate.txt","r+",encoding="utf-8") 
     coordinate=file.readlines() 

    except FileNotFoundError:
      print("coordinate.txt not found!")
      exit()

    file.seek(0) #The top row operation in the coordinate.txt file

     #Determining whether the coordinates are MGRS or Geographic.
    firstline=file.readline()
    spaceindex=firstline.find(" ")
    control=firstline[spaceindex]

    if control == " ":
     print ("Geographic > MGRS conversion type is selected.")
     for c in coordinate:
      
        m = mgrs.MGRS()
        c=c.encode() #Converting string variable to byte data type
        d = m.MGRSToUTM(c)
        d=str(d)
        index=d.find(",")
        l=len(d)
        latitude=d[1:index]
        longitude=d[(index+2):(l-1)]
        file2=open("coordinate_UTM.txt","a+",encoding="utf-8")
        file2.write(longitude+" "+latitude+"\n")

(longitude)

```

Output:
```
'32.852575679878804'
```

```python
convert = mgrs.MGRS()
coords = u"48RYS374671"
convert.MGRSToUTM(coords)
```

Output:
```
(48, 'N', 737400.0, 3167100.0)
```

<!-- grade_id: global_mgrsy_test1, points: 1, grade: true, solution: false -->
```python
'''test your implementation with example test cases'''
assert mgrsToUtmY_Global('48RYS374671')%2000000==1167100, '48RYS374671 should be 1167100 in UTM y coordinate within its 2000Km block'
assert mgrsToUtmY_Global('48RYS374671')==3167100, '48RYS374671 should be 3167100 in UTM y coordinate'
print('success!')
```

<!-- grade_id: cell-e1d0989ec738192e, points: 1, grade: true, solution: false -->
```python
'''hidden tests will be executed after submission'''
```